---
title       : Einheit 4 - inclass
description : ggplot

--- type:NormalExercise lang:r xp:100 skills:1 key:04384eb4ed
## 1. Plot
Sie kennen bereits die plot-Funktion, mit der Sie Datenreihen grafisch darstellen können. Nun wurde in der Vorlesung die ggplot Bibliothek eingeführt. 
Gegeben ist ihnen der Datensatz "diamonds". Darin sind verschiedene Informationen gegeben, wie zum Beispiel Karatzahl, Preis und Qualität. Schauen Sie ihn sich in der Konsole an und erstellen Sie 2 Plots - die Bibliothek ist bereits eingebunden.


*** =instructions
- Erstellen Sie einen Plot der die Karatzahl und den Preis aus dem dataframe "diamonds". Benutzen Sie die plot() Funktion.
- Erstellen Sie den gleichen Plot mit der ggplot Funktion.

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
