---
title       : Einheit 4 - homework
description : Homework zum ggplot Kapitel


--- type:NormalExercise lang:r xp:100 skills:1 key:4262a8c0ff
## ggplot(I)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}
library(ggplot2)

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
