---
title       : Einheit 2 - inclass
description : Arbeiten mit Datensätzen
--- type:NormalExercise lang:r xp:100 skills:1 key:4e6279614f
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
